---
title: 'Chapter 1: Basic'
description: 'In this chapter, we will show you the basic operations in R'
---

## Addition

```yaml
type: VideoExercise
key: 07297dba14
xp: 50
```

`@projector_key`
dc9ec28eb62ddf7df003600fb212eeb1

---

## Exercise 1

```yaml
type: NormalExercise
key: 035b73389d
xp: 100
```

<!-- Basic Operations Test -->
Basic Operations Test

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- Calculate 1+1

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- Just Write 1+1 in Script Section.

`@pre_exercise_code`
```{r}
# Calcuate 1+1
```

`@sample_code`
```{r}
#add print function to print 1+1
1+1
```

`@solution`
```{r}
print(1+1)
```

`@sct`
```{r}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
print(1+1)
```
